<a href="https://colab.research.google.com/github/Gift-eiei/CE880_Case_Study/blob/main/week10/problem1.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Phrugsa Limbunlom (2311569)"
COLLABORATORS = "Phrugsa Limbunlom (2311569)"
```

---

---
# Welcome to CE880
### This is your week-10 : Problem notebook

For this problem set, we'll be using the Jupyter notebook and please upload this notebook to [Google Colab](https://colab.research.google.com/).

Paris Housing is a set of data created from imaginary data of house prices in an urban environment.
https://github.com/sagihaider/CE880_2021/blob/main/Data/ParisHousingClass.csv

All attributes are numeric variables and they are listed bellow:
* squareMeters
* numberOfRooms
* hasYard
* hasPool
* floors - number of floors
* cityCode - zip code
* cityPartRange - the higher the range, the more exclusive the neighbourhood is
* numPrevOwners - number of prevoious owners
* made - year
* isNewBuilt
* hasStormProtector
* basement - basement square meters
* attic - attic square meteres
* garage - garage size
* hasStorageRoom
* hasGuestRoom - number of guest rooms
* price - price of a house
* category - Luxury or Basic

```python
import numpy as np
import pandas as pd
url = 'https://raw.githubusercontent.com/sagihaider/CE880_2021/main/Data/ParisHousingClass.csv'
trainData = pd.read_csv(url,index_col=0)
print(trainData.head())
```

Output:
```
numberOfRooms  hasYard  hasPool  floors  cityCode  \
squareMeters                                                      
75523                     3        0        1      63      9373   
80771                    39        1        1      98     39381   
55712                    58        0        1      19     34457   
32316                    47        0        0       6     27939   
70429                    19        1        1      90     38045   

              cityPartRange  numPrevOwners  made  isNewBuilt  \
squareMeters                                                   
75523                     3              8  2005           0   
80771                     8              6  2015           1   
55712                     6              8  2021           0   
32316                    10              4  2012           0   
70429                     3              7  1990           1   

              hasStormProtector  basement  attic  garage  hasStorageRoom  \
squareMeters       
```

```python
from sklearn.preprocessing import LabelEncoder

def label_encoded(feat):
    le=LabelEncoder()
    le.fit(feat)
    print(feat.name,le.classes_)
    return le.transform(feat)
trainData['category']=label_encoded(trainData['category'])
```

Output:
```
category ['Basic' 'Luxury']
```

```python
from sklearn.model_selection import train_test_split,GridSearchCV

y=trainData['category']
x=trainData.drop('category',axis=1)
xtrain,xtest,ytrain,ytest=train_test_split(x,y,test_size=0.2,random_state=42)
```

---
## Question 1:

#### Classification
Use the dataset, which I have spilted into training and test sets (xtrain, xtest, ytrain, ytest), please train the machine learning model to get the test accuracy of 100%. You are free to use any classication model such as Decision Tree, Random Forest, KNN, SVM..etc.

Hint: Please use grid search to find a best model using different parameters.

```python
from sklearn.metrics import accuracy_score
from sklearn.ensemble import RandomForestClassifier
from sklearn.model_selection import KFold, cross_val_score,GridSearchCV
from sklearn.pipeline import Pipeline
```

<!-- grade_id: my_model_ans, points: 0, grade: false, solution: true -->
```python
def my_model(xtrain,xtest,ytrain,ytest):
    """Write a function to train a machine learning model to get the test accuracy of 100%.
    You are free to use any classication model such as Decision Tree,
    Random Forest, KNN, SVM..etc. """
    # YOUR CODE HERE

    clf = RandomForestClassifier(random_state=42)

    pipe = Pipeline(steps=[('classification', clf)])

    param_grid = {
    'classification__n_estimators': [10, 50, 100, 200],          # Number of trees in the forest
    'classification__max_depth': [None, 10, 20, 30],             # Maximum depth of the tree
    'classification__min_samples_split': [2, 5, 10],             # Minimum number of samples required to split an internal node
    'classification__min_samples_leaf': [1, 2, 4]                # Minimum number of samples required to be at a leaf node
    }

    search_rf_clf = GridSearchCV(pipe, param_grid=param_grid, cv=4)

    search_rf_clf = search_rf_clf.fit(xtrain, ytrain)

    ypred = search_rf_clf.predict(xtest)

    score = accuracy_score(ytest, ypred)

    print(f"Accuracy is {score*100}")

    n_estimators = search_rf_clf.best_params_["classification__n_estimators"]
    max_depth =  search_rf_clf.best_params_["classification__max_depth"]
    min_samples_split =  search_rf_clf.best_params_["classification__min_samples_split"]
    min_samples_leaf =  search_rf_clf.best_params_["classification__min_samples_leaf"]

    print(f"Best parameters: n_estimators = {n_estimators}, max_depth = {max_depth}, min_samples_split = {min_samples_split}, min_samples_leaf = {min_samples_leaf}")

    return score*100

    raise NotImplementedError()


```

<!-- grade_id: my_model_test, points: 9, grade: true, solution: false -->
```python
# Check you solution by running this cell
import math
assert math.isclose(my_model(xtrain,xtest,ytrain,ytest), 100.0, rel_tol = 0.05)
```

Output:
```
Accuracy is 100.0
Best parameters: n_estimators = 50, max_depth = None, min_samples_split = 2, min_samples_leaf = 1
```